Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name below:

```OCaml
let name = "Rishika Varma K"
let rollno = "CS18B045"
```

Output:
```
Findlib has been successfully loaded. Additional directives:
  #require "package";;      to load a package
  #list;;                   to list the available packages
  #camlp4o;;                to load camlp4 (standard syntax)
  #camlp4r;;                to load camlp4 (revised syntax)
  #predicates "p,q,...";;   to set these predicates
  Topfind.reset();;         to force that packages will be reloaded
  #thread;;                 to enable threads

val name : string = "Rishika Varma K"

val rollno : string = "CS18B045"
```

## Important notes about grading:

1. **Compiler errors:** All code you submit must compile. Programs that do not compile will probably receive an automatic zero. If you are having trouble getting your assignment to compile, please visit consulting hours. If you run out of time, it is better to comment out the parts that do not compile, than hand in a more complete file that does not compile.
2. **Late assignments:** Please carefully review the course website's policy on late assignments, as all assignments handed in after the deadline will be considered late. Verify on moodle that you have submitted the correct version, before the deadline. Submitting the incorrect version before the deadline and realizing that you have done so after the deadline will be counted as a late submission.

## Problem 1

Implement predecessor function

```ocaml
pred : int -> int
```

which returns the successor of the given integer.

<!-- grade_id: cell-304066634b90fa89, points: 0, grade: false, solution: true -->
```OCaml
let pred x = x-1
  
  
```

Output:
```
val pred : int -> int = <fun>
```

<!-- grade_id: cell-785fb22b90a8c878, points: 10, grade: true, solution: false -->
```OCaml
assert (pred 1 = 0)
```

Output:
```
- : unit = ()
```

```OCaml

```